```python
# Initialize Otter
import otter
grader = otter.Notebook("lab42.ipynb")
```

**Lab 42: This is a completely made up assignment**

1. Create a function named sum which takes two numbers and sums them

```python
def sum(a,b):
    ...
    
```

```python
grader.check("q1")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Save you notebook first, then run this cell to export your submission as a .zip flie. Upload this file to Gradescope from our Canvas page.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: 1)
```
>>> sum(13,7)
20
```
